---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Introduction to MongoDB - Exercises

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise to practice with MongoDB

`@instructions`
Insert data into Collections and retrieve the information making queries.

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Introduction to MongoDB

```yaml
type: VideoExercise
key: f841653122
xp: 50
```

`@projector_key`
ae71dddb483113c2792334c735bc2434
